# EEN060 - Applied object-oriented programming


Teacher: [Carlos Natalino](https://www.chalmers.se/en/persons/carda/) / Examiner: [Paolo Monti](https://www.chalmers.se/en/persons/mpaolo/)

[Canvas course page](https://chalmers.instructure.com/courses/33383)

[Course channel on Chalmers Play](https://play.chalmers.se/channel/EEN060_EEN065_Applied_object-oriented_programming/300149)

Before you turn this assignment list in, make sure everything runs as expected.
First, **restart the kernel** and then **run all cells**.
Then, check if all the tests run correctly.
Note that if one of the problems present an error, the following ones **will not** be tested.

In case of discrepancies between the problem command and the tests, you should solve it having in mind the tests.

There are two types of cell:
1. *solution cells:* These are the cells where you write your answer, or modify the existing code to solve the problem.
2. *test cells:* These cells are used to test whether your solution is correct or not. If the tests run correctly, you should see a message `tests passed`. Otherwise, you should see an error message.
3. *code quality tests:* These cells tests whether your solution adheres to the code quality rules of the course.

**Delete** the line `raise NotImplementedError()` from the problems that you solve.

**Do not delete or add any cell in this file.** All cells that you need are already in place.

If you want to execute a cell, select the cell and press **CTRL+Enter** (in Windows) or **CMD+Enter** (in macOS) or click on the **Run cell** button.


## Assignment Week 3

In this assignment, students should practice:
- how to use Visual Studio Code
- how to download, solve and submit programming assignments
- how to create functions
- start getting familiar with Python
- how to manipulate variables
- how to use math operations
- how to output the results
- understand different types of variables
- understand the code quality validation

**Observations:**
For the problems in this assignment, you should:
- use only the math operations studied in the lecture notes (i.e., no advanced functions allowed).
- not use external modules, i.e., `import` command is not allowed in your solutions (with exception of the module `typing`).
- make sure that your `utils.py` file has the following line (i.e., replace `onu2` by `onu1` if needed):

<!-- **You should need only Chapters 1-3 of the lecture notes to solve all the problems in this list.** -->

**Preparation:** Run the cell below every time you start working on this file, and every time you restart the kernel.

```python
%load_ext autoreload
# check if this file is within the same folder as the `utils.py` file.
try:
    from utils import validate_python_code
except:
    import sys
    print("It seems this file is in the wrong folder. "
          "Make sure to place it in the `programming-assignments` folder/project.",
          file=sys.stderr)
```

```python
from utils import validate_python_code
```

```python
import datetime
import getpass
import os
import sys

# check if the virtual env is correct
venv_path = os.path.abspath(sys.executable)
cur_path = os.path.abspath(os.getcwd())
assert os.path.commonpath([venv_path, cur_path]) == cur_path, \
    "It seems like this file is being executed with the wrong virtual environment."

print("Version:", sys.version)
print("Executable (virtual env):", sys.executable)
print("Current folder:", os.getcwd())
print("Datetime:", datetime.datetime.now())
print("User:", getpass.getuser())
del datetime, getpass, os, sys
```

Output:
```
Version: 3.13.1 (main, Dec  3 2024, 17:59:52) [Clang 16.0.0 (clang-1600.0.26.4)]
Executable (virtual env): /Users/carda/projects/programming-assignments-2025/venv/bin/python
Current folder: /Users/carda/projects/programming-assignments-2025
Datetime: 2025-02-03 11:52:03.593045
User: carda
```

### Problem 1

Create a function called `assignment_points` that receives an integer representing the number of assignment approved.
The function then returns a float that represents the assignment points obtained according to the scale defined in the [course PM](https://chalmers.instructure.com/courses/33383#:~:text=1%20AP,15%20APs).

Make sure to return `0` if no assignment is approved, i.e., when you recieve the input `0`.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

```python
%%writefile assignment_points_solution.py
# solution cell
### BEGIN SOLUTION
raise NotImplementedError()
### END SOLUTION
```

```python
%autoreload 2
# uncomment the line(s) below to debug
# from assignment_points_solution import assignment_points

# count = 6
# ap = assignment_points(count)
# print(f"The assignment points are {ap}")
# print("Execution finished", "\u2713")
```

```python
%autoreload 2
# test cell
try:
    import assignment_points_solution
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from assignment_points_solution import assignment_points
except:
    raise ValueError("Your solution does not contain the right function!")

test_cases = [(6, 11.0), \
        (7, 15.0), \
        (4, 6.0), \
        (1, 1.0), \
        (0, 0.0)]

for _in, _out in test_cases:
    _res = assignment_points(_in)
    assert _res == _out, f"The function with input `{_in}` should return the value \
`{_out}` of type `{type(_out)}`\n but returned the value `{_res}` of type `{type(_res)}`."

print("tests passed", "\u2713")
```

```python
validate_python_code("assignment_points_solution.py")
```

### Problem 2

Create a function called `smallest_value` that receives three numbers (`int` or `float`) as an input, and returns the smallest one among them.

**Note:**
- You must not use built-in functions to solve the problem; use only conditional statements.
- You must not use any import statement.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

```python
%%writefile smallest_value_solution.py
# solution cell
### BEGIN SOLUTION
raise NotImplementedError()
### END SOLUTION
```

```python
%autoreload 2
# uncomment the line(s) below to debug
# from smallest_value_solution import smallest_value

# a, b, c = 12, 12, 11
# l = smallest_value(a, b, c)
# print(f"The smallest value among {a}, {b}, {c} is {l}")
# print("Execution finished", "\u2713")
```

<!-- grade_id: cell-test-1-2, points: 1, grade: true, solution: false -->
```python
%autoreload 2
# test cell
try:
    import smallest_value_solution
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from smallest_value_solution import smallest_value
except:
    raise ValueError("Your solution does not contain the right function!")

test_cases = [((1, 2.2, 3), 1), ((2, 1.1, 3), 1.1), ((3, 2.4, 2.3), 2.3), \
    ((1, -2, -7), -7)]

for _in, _out in test_cases:
    _res = smallest_value(*_in)
    assert _res == _out, f"The function with input `{_in}` should return the value \
`{_out}` of type `{type(_out)}`\n but returned the value `{_res}` of type `{type(_res)}`."

print("tests passed", "\u2713")
```

```python
validate_python_code("smallest_value_solution.py")
print("code quality passed", "\u2713")
```

### Problem 3

Create a Python function named `reverse_lower_and_upper_case` that receives a string as an input. It reverses the case of the input. That is if the input is in upper case, then convert it to lower case and if it is in lower case, then convert it to lower case.

**Examples**

Input: 'test'

Output: 'TEST'

Explanation: Here 'test' is in lower case, so we converted it to upper case.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

```python
%%writefile reverse_lower_and_upper_case_solution.py
# solution cell
### BEGIN SOLUTION
raise NotImplementedError()
### END SOLUTION
```

```python
%autoreload 2
# uncomment the line(s) below to debug
from reverse_lower_and_upper_case_solution import reverse_lower_and_upper_case

initial = ['test', 'PRACTICE']
for _i in initial:
    result = reverse_lower_and_upper_case(_i)
    print(f"The reverse case version of {_i} is {result}")
print("Execution finished", "\u2713")
```

```python
%autoreload 2
# test cell
try:
    import reverse_lower_and_upper_case_solution
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from reverse_lower_and_upper_case_solution import reverse_lower_and_upper_case
except:
    raise ValueError("Your solution does not contain the right function!")

test_cases = [('test', 'TEST'), \
    ('practice', 'PRACTICE'), \
    ('PRACTICE', 'practice'), \
    ('RESOURCE NOT GRANTED', 'resource not granted')]

for _in, _out in test_cases:
    _res = reverse_lower_and_upper_case(_in)
    assert _res == _out, f"The function should return the value `{_out}` of type \
`{type(_out)}`\n but returned the value `{_res}` of type `{type(_res)}`."
print("tests passed", "\u2713")
```

```python
validate_python_code("reverse_lower_and_upper_case_solution.py")
print("code quality passed", "\u2713")
```

### Problem 4

Create a Python function called `mobile_number_validate` that receives a mobile number as a string and validates whether it is a correct mobile number (returning `True`) or not (returning `False`). 

Rules for mobile number:
- string starts with either of `0`, `46`, or `+46`;
- the remaining string has length 9 such that the first digit is between 1 to 9 and the other digits are between 0 to 9.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

```python
%%writefile mobile_number_validate_solution.py
# solution cell
### BEGIN SOLUTION
raise NotImplementedError()
### END SOLUTION
```

```python
%autoreload 2
# uncomment the line(s) below to debug
# from mobile_number_validate_solution import mobile_number_validate

# number = "46102345678" # Valid
# valid = mobile_number_validate(number)
# valid_string = "valid" if valid else "invalid"
# print(f"The mobile number {number} is {valid_string}")
# print("Execution finished", "\u2713")
```

```python
%autoreload 2
# test cell
try:
    import mobile_number_validate_solution
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from mobile_number_validate_solution import mobile_number_validate
except:
    raise ValueError("Your solution does not contain the right function!")

with open("mobile_number_validate_solution.py", "rt", encoding="UTF-8") as file:
    for idl, line in enumerate(file.readlines()):
        line = line.replace("\n", "")
        assert "import" not in line, f"You must not use any `import` statement. Found one in line {idl+1} with content `{line}`."

test_cases = [('0061234567', False), ('+46061234567', False), ('123456789', False), \
    ('+465555', False), ('46123456789', True), ('+46761234567', True)]
for _in, _out in test_cases:
    _res = mobile_number_validate(_in)
    assert _res == _out, f"The function should return the value `{_out}` of type \
`{type(_out)}`\n but returned the value `{_res}` of type `{type(_res)}`."

print('tests passed', u'\u2713')
```

```python
validate_python_code("mobile_number_validate_solution.py")
print('tests passed', u'\u2713')
```