# EEN060 - Applied object-oriented programming

Teacher: [Carlos Natalino](https://www.chalmers.se/en/persons/carda/) / Examiner: [Paolo Monti](https://www.chalmers.se/en/persons/mpaolo/)

[Canvas course page](https://chalmers.instructure.com/courses/33383)

[Course channel on Chalmers Play](https://play.chalmers.se/channel/EEN060_EEN065_Applied_object-oriented_programming/300149)

Before you turn this assignment list in, make sure everything runs as expected.
First, **restart the kernel** and then **run all cells**.
Then, check if all the tests run correctly.
Note that if one of the problems present an error, the following ones **will not** be tested.

In case of discrepancies between the problem command and the tests, you should solve it having in mind the tests.

There are two types of cell:
1. *solution cells:* These are the cells where you write your answer, or modify the existing code to solve the problem.
2. *test cells:* These cells are used to test whether your solution is correct or not. If the tests run correctly, you should see a message `tests passed`. Otherwise, you should see an error message.

**Do not delete or add any cell in this file.** All cells that you need are already in place.

<!-- You may also have problems that will be manually graded (e.g., HTML/CSS/JS/pseudocode problems).
These problems **do not** have a respective *test cell*. -->

If you want to execute a cell, select the cell and press **CTRL+Enter** (in Windows) or **CMD+Enter** (in macOS) or click on the **Run cell** button.


# Lab assignment


```python
%load_ext autoreload
try:
    from utils import validate_python_code
except:
    import sys
    print("It seems this file is in the wrong folder. "
          "Make sure to place it in the `programming-assignments` folder/project.",
          file=sys.stderr)
```

Output:
```
The autoreload extension is already loaded. To reload it, use:
  %reload_ext autoreload
```

### Problem 1

Create a Python function called `calculate_addition` that receives a two positive integers $a$ and $b$ and returns the sum $a+b$.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-1-1, points: 0, grade: false, solution: true -->
```python
%%writefile calculate_addition_solution.py
# solution cell


def calculate_addition(a: int, b: int) -> int:
    c = a + b
    return c

```

Output:
```
Overwriting calculate_addition_solution.py
```

```python
%autoreload 2
# uncomment the line(s) below to debug
from calculate_addition_solution import calculate_addition


a = 7
b = 2
c = calculate_addition(a, b)
print(f"The sum of a = {a} and b = {b} is {c}")
print("Execution finished", "\u2713")

```

Output:
```
The sum of a = 7 and b = 2 is 9
Execution finished ✓
```

<!-- grade_id: cell-tests-1-1, points: 1, grade: true, solution: false -->
```python
%autoreload 2
# test cell
try:
    import calculate_addition_solution
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from calculate_addition_solution import calculate_addition
except:
    raise ValueError("The module does not have the necessary function!")

test_cases = [((2, 4), 6), ((3, 9), 12), ((7, 2), 9)]

for _in, _out in test_cases:
    _res = calculate_addition(*_in)
    assert _res == _out, f"The function with input `{_in}` should return the value \
`{_out}` of type `{type(_out)}`\n but returned the value `{_res}` of type `{type(_res)}`."
```

```python
validate_python_code("calculate_addition_solution.py")
print("tests passed", "\u2713")
```

Output:
```
<IPython.core.display.HTML object>
tests passed ✓
```

### Problem 2

Create a Python function called `calculate_subtraction` that receives a two positive integers $a$ and $b$ and returns the difference $a-b$.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

```python
%%writefile calculate_subtraction_solution.py
# solution cell


def calculate_subtraction(a: int, b: int) -> int:
    c = a - b
    return c

```

Output:
```
Overwriting calculate_subtraction_solution.py
```

```python
%autoreload 2
# uncomment the line(s) below to debug
from calculate_subtraction_solution import calculate_subtraction


a = 7
b = 2
c = calculate_subtraction(a, b)
print(f"The sum of a = {a} and b = {b} is {c}")
print("Execution finished", "\u2713")

```

Output:
```
The sum of a = 7 and b = 2 is 5
Execution finished ✓
```

```python
%autoreload 2
# test cell
try:
    import calculate_subtraction_solution
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from calculate_subtraction_solution import calculate_subtraction
except:
    raise ValueError("The module does not have the necessary function!")

test_cases = [((2, 4), -2), ((3, 9), -6), ((7, 2), 5)]

for _in, _out in test_cases:
    _res = calculate_subtraction(*_in)
    assert _res == _out, f"The function with input `{_in}` should return the value \
`{_out}` of type `{type(_out)}`\n but returned the value `{_res}` of type `{type(_res)}`."
```

```python
validate_python_code("calculate_subtraction_solution.py")
print("tests passed", "\u2713")
```

Output:
```
<IPython.core.display.HTML object>
tests passed ✓
```

### Problem 3

Create a Python function called `calculate_product` that receives a two positive integers $a$ and $b$ and returns the product $ab$.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

```python
%%writefile calculate_product_solution.py
# solution cell


def calculate_product(a: int, b: int) -> int:
    c = a * b
    return c

```

Output:
```
Overwriting calculate_product_solution.py
```

```python
%autoreload 2
# uncomment the line(s) below to debug
from calculate_product_solution import calculate_product


a = 7
b = 2
c = calculate_product(a, b)
print(f"The sum of a = {a} and b = {b} is {c}")
print("Execution finished", "\u2713")

```

Output:
```
The sum of a = 7 and b = 2 is 14
Execution finished ✓
```

```python
%autoreload 2
# test cell
try:
    import calculate_product_solution
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from calculate_product_solution import calculate_product
except:
    raise ValueError("The module does not have the necessary function!")

test_cases = [((2, 4), 8), ((3, 9), 27), ((7, 2), 14)]

for _in, _out in test_cases:
    _res = calculate_product(*_in)
    assert _res == _out, f"The function with input `{_in}` should return the value \
`{_out}` of type `{type(_out)}`\n but returned the value `{_res}` of type `{type(_res)}`."
```

```python
validate_python_code("calculate_product_solution.py")
print("tests passed", "\u2713")
```

Output:
```
<IPython.core.display.HTML object>
tests passed ✓
```

### Problem 4

Create a Python function called `calculate_division` that receives a two positive integers $a$ and $b$ and returns the quotient $a/b$.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

```python
%%writefile calculate_division_solution.py
# solution cell


def calculate_division(a: int, b: int) -> float:
    c = a / b
    return c

```

Output:
```
Overwriting calculate_division_solution.py
```

```python
%autoreload 2
# uncomment the line(s) below to debug
from calculate_division_solution import calculate_division


a = 7
b = 2
c = calculate_division(a, b)
print(f"The sum of a = {a} and b = {b} is {c}")
print("Execution finished", "\u2713")

```

Output:
```
The sum of a = 7 and b = 2 is 3.5
Execution finished ✓
```

```python
%autoreload 2
# test cell
try:
    import calculate_division_solution
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from calculate_division_solution import calculate_division
except:
    raise ValueError("The module does not have the necessary function!")

test_cases = [((2, 4), 0.5), ((3, 9), 0.3333333333333333), ((7, 2), 3.5)]

for _in, _out in test_cases:
    _res = calculate_division(*_in)
    assert _res == _out, f"The function with input `{_in}` should return the value \
`{_out}` of type `{type(_out)}`\n but returned the value `{_res}` of type `{type(_res)}`."
```

```python
validate_python_code("calculate_division_solution.py")
print("tests passed", "\u2713")
```

Output:
```
<IPython.core.display.HTML object>
tests passed ✓
```